---
title: Test
description: >-
  Test


---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: 81b8736e0a
```

Ceci est l'assignment (Describes the task the student is to solve).

```
var example = "hello!";
alert(example);
```

`@instructions`
Ceci est l'instruction (The actual question or task for student).

`@hint`
Ceci est le hint (A hint student can request).

`@pre_exercise_code`
```{python}
test = 2
```
`@sample_code`
```{python}
print(3+2)
```
`@solution`
```{python}
print(3*2)
```
`@sct`
```{python}
test_student_typed("abc")
success_msg('Great job!')
```


